SIV-specific neutralizing antibody induction following selection of a PI3K drive-attenuated nef variant

Author response:

The following is the authors’ response to the original reviews.

Reviewer #1 (Public Review):

This work describes the induction of SIV-specific NAb responses in rhesus macaques infected with SIVmac239, a neutralization-resistant virus. Typically, host NAb responses are not detected in animals infected with SIVmac239. In this work, seventy SIVmac239-infected macaques were retrospectively screened for NAb responses and a subset of nine animals were identified as NAb-inducers. The viral genomes from 7/9 animals that induced NAb responses were found to encode nonsynonymous mutation in the Nef gene (amino acid G63E). In contrast, Nef G63E mutation was found only in 2/19 NAb non-inducers - implicating that the Nef G63E mutation is selected in NAb inducers. Measurement of Nef G63E frequencies in plasma viruses suggested that Nef G63E selection preceded NAb induction. Nef G63E mutation was found to mediate escape from Nef-specific CD8+ T-cell responses. To examine the functional phenotype of Nef G63E mutant, its effect on downmodulation of Nef-interacting host proteins was examined. Infection of rhesus and cynomolgus macaque CD4+ T cell lines with WT or Nef G63E mutant SIV suggested that Nef mutant reduces S473 phosphorylation of AKT. Using flow cytometry-based proximity ligation assay, it was shown that Nef G63E mutation reduced binding of Nef to PI3K p85/p110 and mTORC2 GβL/mLST8 and MTOR components - kinase complex responsible AKT-S473 phosphorylation. In vitro B-cell Nef invasion and in vivo imaging/flow cytometry-based assays were employed to suggest that Nef from infected cells can target Env-specific B cells. Lastly, it was determined that NAb inducers have significantly higher Env-specific B-cells responses after Nef G63E selection when compared to NAb non-inducers. Finally, a corollary was drawn between the Nef G63E-associated B-cell/NAb induction phenotype and activated PI3K delta syndrome (APDS), which is caused by activating GOF mutations in PI3K, to suggest that Nef G63E-meidated induction of NAb response is reciprocal to APDS.

Strengths:

This study aims to understand the viral-host interaction that governs NAb induction in SIVmac239-infected macaques - this could enable identification of determinants important for induction of NAb responses against hard-to-neutralize tier-2/3 HIV variants. The finding that SIV-specific B-cell responses are induced following Nef G63E CD8+ T-cell escape mutant selection argue for an evolutionary trade-off between CTL escape and NAb induction. Exploitation of such a cellular-humoral immune axis could be important for HIV/AIDS vaccine efforts.

Although more validation and mechanistic basis are needed, the corollary between PI3K hyperactive signaling during autoimmune disorders and Nef-mediated abrogated PI3K signaling could help identify novel targets and modalities for targeting immune disorders and viral infections.

We are grateful for the supportive and insightful comments. The work did seem to unintendedly highlight a conceptual link between extrinsic and intrinsic immune perturbations. We will keep working on both wings, aiming to evoke synergisms.

Weaknesses:

Although the paper does have strengths in principle, the weaknesses of the paper are that the mechanistic basis of Nef-mediated induction of NAb responses are not directly examined. For example, it remains unclear whether SIVmac239 with engineered G63E mutation in Nef would induce faster and potent NAb responses. A macaque challenge study is needed to address this point.

We appreciate the point. We do have certain difficulties in availability of macaques for de novo experiments. As partially discussed in ver1, the identified Nef phenotype selected post-acute infection confers an enhanced CD4+ T cell-killing effect (revised Fig 4F), and it is likely that de novo infection with the mutant would redirect the trajectory of infection to rapid disease/AIDS progression accompanying generalized immune failure by boosting acute-phase CD4 destruction. In other words, mutant de novo infection may not necessarily be directly discussable as an attempt for reconstitution. It appears equally critical to understand the mutant in vitro on an immunosignaling basis, and in the current work we have focused on depicting this as the first step. We will work on reconstitution experiments with emphasis on pharmacology in our future study.

As presented, the central premise of the paper involves infected cell-generated Nef (WT or G63E mutant) being targeted to adjacent Env-specific B cells. However, it remains unclear how this is transfer takes place. A direct evidence demonstrating CD4+ T cell-associated and/or cell-free Nef being transferred to B-cell is needed to address this concern.

We appreciate the point, also pointed out by Reviewers 2 and 3. We have performed three sets of in vitro reconstitution experiments graphically/functionally addressing how Nef transfer from CD4+ T cells to B cells can be modulated (new Fig 6) and edited text accordingly.

The interaction between Nef and PI3K signaling components (p85, p110, GβL/mLST8, and MTOR) has been explored using PLA assay, however, this requires validation using additional biochemical and/or immunoprecipitation-based approaches. For example, is Nef (WT or mutant form) sufficient to affect PI3K-induced phosphorylation of Akt in an in vitro kinase assay? Moreover, the details regarding the binding events of WT vs mutant Nef with PI3K signaling components is lacking in this study. Lastly, it is unclear whether the interaction of Nef with PI3K signaling components is a conserved function of all primate lentiviruses or is this SIV-specific phenotype.

We appreciate the point. Co-immunoprecipitation analysis via pulldown with the mTORC2-intrinsic cofactor Sin1 (revised Fig 4E), showing decreased G63E-Nef binding, should confer robustness to the statement combined with initial manipulation results (Fig 4C). As Sin1 is mTORC2- and not mTORC1-intrinsic, results should be strengthened. Phosflow may be a standard readout nowadays for pAkt itself. Related with sequence variation, conservation will be addressed in studies ahead. We concisely mentioned on this in the revision (Lines 390-391).

It has been previously reported that the region of Nef encoding glycine at position 63 is not conserved in HIV-1 (Schindler et al, Journal of Virology 2004). Thus, does HIV-1 Nef also function in induction of NAb responses in humans? or the observed phenotype specific to SIV?

We appreciate the point, and do not have an answer at the moment. We will explore in our HIV-1-infected patient cohort (Hau et al, AIDS 2022) and other occasions whether corresponding phenotypes may exist. We have mentioned on this point in the revised manuscript (Line 392-393).

Reviewer #2 (Public Review):

It is well known that human and simian immunodeficiency viruses (HIV and SIV, respectively) evolved numerous mechanisms to compromise effective immune responses but the underlying mechanisms remain incompletely understood. Here, Yamamoto and Matano examined the humoral immune response in a large number of rhesus macaques infected with the difficult-to-neutralize SIVmac239 strain. They identified a subgroup of animals that showed significant neutralizing Ab responses. Sequence analyses revealed that in most of these animals (7/9) but only a minority in the control group (2/19) SIVmac variants containing a CD8+ T-cell escape mutation of G63E/R in the viral Nef gene emerged. They further show that this change attenuates the ability of Nef to stimulate PI3K/Akt/mTORC2 signaling. The authors propose that this induction of SIVmac239 nAb induction is reciprocal to antibody dysregulation caused by a previously identified human PI3K gain-of-function (Ref). Altogether, the results suggest that PI3K signaling plays a key role in B-cell maturation and generation of effective nAb responses.

Strengths of the study are that the authors analyzed a large number of SIVmac-infected macaques to unravel the biological significance of the known effect of the interaction of Nef with PI3K/Akt/mTORC2 signaling. This is interesting and may provide a novel means to improve humoral immune responses to HIV. Weaknesses are that only G63E and not G63R that also emerged in most animals was examined in most functional assays. Some effects of the G63E mutation seem modest and comparison to a grossly nef-defective SIVmac construct would be desirable to better assess to impact of the mutation of Nef-mediated stimulation of PI3K. While the impact of this Nef mutations on PI3K and the association with improved nAb responses is largely convincing, the results on the potential impact of soluble Nef on neighboring B cells is much less clear. SIVmac239 infects and manipulates helper CD4 T cells and these are essential for the activation and differentiation of B cells into antibody-producing plasma cells and effective humoral immune responses. Without additional functional evidence that Nef indeed specifically targets and manipulated B cells these results and conclusions should be made with much greater caution. Finally, the presentation of the results and conclusions is partly very convoluted and difficult to comprehend. Editing to improve clarity is highly recommended.

We are very grateful for the supportive and visionary review and suggestions. Experiments have been performed to improve the points raised. This work inevitably involved interdisciplinary factors to even hit on the schematic (NAbs, B cells, CD4+T, CD8+T, viral escape, immunosignaling, IEI as extrapolation & microscopy implementations) and convoluted sections should have existed. We attempted streamlining of certain portions and edited writing throughout, and hope that it became more straightforward.

Reviewer #2 (Recommendations For The Authors):

As outlined in the public review, I found the results potentially very interesting but parts of the manuscript much more complex and confusing than necessary. In addition, the methods on the potential impact of soluble Nef on neighboring B cells in vivo was difficult to assess but altogether this part was not convincing. Have the following specific suggestions:

We are very grateful for the scholarly review, and encouraging and suggestive comments on this orphan work. In the revision we designed experiments to address the properties of Nef transfer to append understanding on the in vivo B-cell data. Recommendations have been addressed as follows.

(1) Title: "AIDS virus-neutralizing antibody induction reciprocal to a PI3K gain-of-function disease". Think this title hardly reflects the data; SIVmac cause simian AIDS and is not the "AIDS virus" the 2nd part is more appropriate for discussion than for the title (and the abstract).

We appreciate the point. The original intent of the title was to conceptually bridge two differing fields of virus-host interaction and inborn errors of immunity/immunosignaling on an original article basis. Certain papers (Mudd et al, Nature 2012 etc) do utilize the term AIDS virus, and we similarly chose the term for simplification to non-virologists at initial submission.

That being said, we understand the scholarly point raised, and feel that the initial aim can be well attained by retaining the key host effector PI3K in the title, as in the revised submission titled “SIV-specific neutralizing antibody induction following selection of a PI3K drive-attenuated nef variant”.

(2) Abstract and throughout: As the authors show, SIVmac is not generally "neutralization resistant"; difficult to neutralize is more appropriate and should be used throughout. Also, the abstract and other parts are more complicated than necessary.

We appreciate the point. HIV/SIV Env immunology work utilizes “neutralization-resistant” for SIVmac239 (e.g., Mason et al, PLoS Pathog 2016), and autologous titer positivity of ~10% at this size of examination does appear low amongst lentiviruses. Nevertheless, as recommended, “difficult-to-neutralize” better describes the nature, and we have switched the term accordingly.

Linked with title modification, we reflected the comment on abstract structure and switched the main introductory sentence (Here we…) to a more data-based one instead of depicting extrapolation, and have modified phrasings in the latter half.

(3) The intro seems a bit biased. Immune evasion due to mutations and proviral integration that play key roles in viral persistence are not mentioned. nAbs are not known to efficiently control HIV or SIV replication in vivo (not even in the present study). Thus, a more "balanced" presentation of the role of nAbs in vivo is desirable.

We agree with the comment. Introduction in ver1 submission was compressed to just display humoral immune perturbation examples across persistence-prone viral infections, and indeed it should be much better to layout the multiscale strategies of lentiviruses in manifesting viral persistence. We have appended two sets of texts, one on the fundamental integrating retroviral life cycle and another on the wide spectrum of accessory protein-driven perturbation. As pointed out, the current endogenous induction is of course not early enough to exert suppressive impact on replication as like in exogenous Ab passive infusions. We have accordingly modulated text to improve the balance.

(4) Lines 73-76: rephrase for clarity.

We acknowledge the comment and have rephrased accordingly.

(5) Line 92: "linked with sustained Env-specific B-cell responses after the mutant Nef selection". After or during in one case; the time frame varies enormously and this should be discussed.

We appreciate the comment. The six Nef-G63E mutant-selecting NAb inducers subjected to B-cell analysis were the ones that showed precedence in Fig 2D (mutant before induction). That being said, we modified text as suggested (Line 104 in revised uploaded text). Text related to temporal deviation has been appended (Lines 378-383 in revised uploaded text).

(6) The authors should discuss G63R and include it in the functional analyses.

We appreciate the comment. Discussion on Nef-G63R in ver1 submission was kept minimal because statistical significance for selection was marginal. We generated a Nef-G63R mutant and results are appended in Fig 4-Figure Supplement 2.

(7) Lines 124/5: conservation only applies to SIVsmm/mac Nefs and this region is also frequently deleted/length-variable in primary HIV-1 Nefs.

We appreciate the comment. We modified description of the region accordingly (Lines 139-141 in revised text).

(8) Lines 153-155: Statement doesn't seem to make sense. The triple mutant Nef SIVmac construct was not attenuated for replication but specifically disrupted in CD3 down-modulation.

We acknowledge the comment. It had meant that the consequent plasma viral load showed a trend of decrease (as in the Graphical Abstract of the work) which should (in a simplistic view) influence antigenicity for humoral immune responses. Yet it is very true that virological replicative capacity was comparable with wild-type as in Fig.1. We have taken down the related text and rephrased it (Ref remains cited in introduction).

(9) Lines 178/9: levels in PI3K gain-of-function mice "with full disease phenotype (Avery et al., 2018)". This needs more information, e.g. what disease exactly are they talking about?

We are grateful for the correction, and have appended text and introduced the mentioned congenital disease in the Introduction section in advance. In-detail description is also appended in the Discussion section.

(10) Lines 186/7: "Env-stimulating high-MOI infection also accelerated phenotype appearance, with enhanced 50% reduction (Figure 4C, right)". Modify text and corresponding figure for clarity.

We acknowledge the comment. We revised as: “A high-MOI SIV infection, comprising higher initial concentration of extracellular Env stimuli, also accelerated phenotype appearance from day 3 to day 1 post-infection with stronger pAkt reduction”.

(11) The validity of the results described in the section "Targeting of lymph node Env-specific B cells by Nef in vivo" was difficult to assess. Altogether, however, I didn't find them convincing, especially since a negative control (e.g. macaques infected with nef-deleted SIVmac) are missing.

We acknowledge the comment. As a pure experimental control, whole-Nef deletion may assist for subtracted baselines. Within this work, the staining per se at least should be highly specific (mAb multiply verified in other applications and cytometry panel also designed for minimal spillover into AF488 channel). On in vivo basis, direct comparison may be somewhat frustrated by the fact that reduction in other pleiotropic effects of Nef seem to more dominate upon Nef deletion, as a set of reduced viremia, robust CD8 responses, killer CD4 responses and increased binding Ab titers (Johnson et al, J Virol 1997, Gauduin et al, J Exp Med 2006, Fukazawa et al, Nat Med 2012, Adnan et al, PLoS Pathog 2016 etc) leading to altered trajectory. We promise that we will work on refinement of the methodology in studies ahead.

(12) Lines 309-319: This paragraph made little sense to me (as did lines 328-331).

We acknowledge the comment and have edited both sections.

Reviewer #3 (Additional Reviewer):

In this manuscript, Hiroyuki Yamamoto et al examined virus-specific antibody responses and identified a subgroup of nine individuals, out of seventy SIVmac239 rhesus macaques of Burmese origin infected with SIVmac239, that develop neutralizing antibodies (NAb). The authors propose the emergence of a nef mutant (Nef-G63E) that impacts on B cell maturation resulting in PI3K gain-of-function.

My major concerns are:

The authors by different aspect addressed the role of the emergence of Nef-G63E mutant in individuals developing NAb. The manuscript is confused and the rational not always clearly stated. This reflects the two aspects of the manuscript (i) NAb identification in a subgroup of macaque and (2) the identification this nef mutation.

We are grateful for the comprehensive and scholarly comments. As pointed out, the work did need to confront potential bifurcation of the influence of the obtained viral immunosignaling phenotype for CD4-intrinsic (which might be your specialty) and B-cell-intrinsic impact. Based on your suggestions we have acquired additional data and revised the manuscript as attached.

The authors used both males (n=57) and females (n=13). However, there is no indication related to the sex regarding NAb inducers versus non-NAb Inducers. The notion of "highly pathogenic" is certainly not correct (see the introduction). Pathogenicity is also depending on monkey origin. Thus, cynomolgus are less sensitive to SIVmac239 or SIVmac251 compared to rhesus macaques (Ling B Aids 2002; Reimann KA, J Virol 2005; Cumont MC, J Virol 2008), or to pigtails used in US. Indeed, the authors used Burmese macaques, and therefore the dynamics of pathogenicity is different to rhesus macaque (Indian origin) housed in US. How many animals have been sacrificed out of the 61 animals? Herein, the animals are surviving longer (more than one year), and therefore the notion of "highly pathogenic" merits to be modulated.

We appreciate the comment. We have accordingly appended sex information (M/F: 8/1 versus 49/12 in NAb inducers vs non-inducers, p > 0.99 by Fisher’s exact test) in the methods section. As pointed out there are differences in the frequency and rate of AIDS progression among macaques of differing origin, whereas we have also previously reported reproducible AIDS progression dependent on MHC-I genotypes in the Burmese rhesus macaques utilized (Nomura, Yamamoto et al., J Virol 2012). Adhering to advice, we have attenuated the term to “pathogenic” in the revised manuscript and appended one reference showing pathogenesis gradation from a cell-death perspective (Cumont 2008).

Furthermore, no indication is provided regarding CD4 T cell dynamics, or CD8 T cells. In particular, the extent of T cell immunodeficiency may compromise humoral response. Therefore, this data needs to be shown. Indeed, previous reports have indicated that early CD4 T cell depletion is associated with defective humoral response. Furthermore, Tfh cell depletion was reported in several immune tissues, which are essential for B cell immune response like the spleen. Thus, this should be discussed as an alternative mechanism to the absence of NAb. Indeed, the authors found higher and persistent env-specific plasmablast cells in NAb inducers than that observed in non-NAb inducers figure 6. Why to have selected twelve individuals out of 61 individuals for assessing anti-env response (Supplemental S3 for figure 1, panel 1), and only eleven for western blots. The explanation in the text is absent. This requires to be clearly stated. See lines 108-110.

We appreciate the comment. As in other sections, this study utilized available cryopreserved samples from a retrospective cohort, also having heterogeneity in data acquisition along the way. We acknowledge that some supplemental data are particularly limited in information, which is also a reason they are presented in SI. We felt that one important core was to secure samples for Nef-G63E-selecting NAb inducers versus viremic non-inducers, for which we acquired six versus twelve in the B-cell analysis.

We (Nakane et al, PLoS ONE 2013) and others (Hirsch et al, J Virol 2004) have already reported on western blotting-basis that SIV-infected rapid progressors tend to manifest serological failure (impaired binding Ab-WB bands). Therefore, to compare quantitative traits at this basal stage (Fig 1), we judged that NAb inducer comparison with more non-rapid-progressing (>60 wk survival) non-inducers would be a criterion. We have mentioned on this in the revised manuscript (results/methods). Additionally, we have replaced the immunoblotting result with one more non-inducer (n = 12) to enhance results. Please note that there are lot deviations in strip-coated antigen (e.g., gp160) but the result is comparable (now covers 12/13 of animals with >60-wk survival).

The authors indicated the frequencies of Nef-G63E mutant in figure 2 panel C. However. no information is indicated in the legend about the number of NAb non-inducers used to calculate this frequency. The authors indicated line 127, "only in two of the nineteen NAb non-inducers, including one rapid progressor". Thus, different numbers of individuals are used through the manuscript. For the readers, this is clearly a statement that needs to be clarify and to refer to what. This is not homogeneous along the text and the analyses performed.

We appreciate the comment, and have appended the number in the revised Fig 2C. As aforementioned, heterogeneity of sample number in different sections is indeed a limitation of the work, and have mentioned this in the Discussion.

The rational related to the sentence lines 140-142. Please clarify.. "NAb induction is not associated with these MHC-I genotypes (P = 0.25 by Fisher's exact test, data not shown) but with the Nef-G63E mutation itself".

We appreciate the comment. We have rephrased it as:

“Ten of nineteen NAb non-inducers also had either of these alleles (Figure 1-figure supplement 1). This did not significantly differ with the NAb inducer group (P = 0.25 by Fisher’s exact test, data not shown), indicating that NAb induction was not simply linked with possession of these MHC-I genotypes but instead required furthermore specific selection of the Nef-G63E mutation.” (Lines 159-162).

In supplemental figure 3, only 7 individuals have been tested, while the authors indicated "Ten of nineteen NAb non-inducers also had either of these alleles". Why only seven? In NAb Burmese monkeys, the authors indicate specific T cells capable to recognize WT nef peptide, but not G63E peptide mutant. Thus, nef is immunogenic in vivo generating T cells despite to be mutated.

In contrary, non-NAb-inducers demonstrate the absence of nef specific T cells (supplemental figure 3, excepted R01-011 panel A). Although, the authors propose an escape mutant for CD8 T cells, this is not associated with the absence of immunogenicity and not with a difference in viral load in comparison to NAb inducers (panel C). Therefore, the conclusions merit to be revised. Thus, this part of the manuscript is confusing. Please clarify the rational to link NAb and Nef specific CD8 T cells.

We appreciate the comment. 7 out of 8 non-inducers positive for the allele and not selecting for the Nef-G63E mutant was available for analysis. The relative contribution of this single Nef62-70 epitope-specific CTL response is speculated not to be largely impacting viral control, among the many induced. This is basally discussed in a previous paper (Nomura, Yamamoto et al., J Virol 2012), more suggestive of an MHC-I haplotype-level correlation with plasma viral load. We assume that the CTL pressure-driven selection of Nef-G63E mutant was a rather pure immunosignaling trigger under persistent viremia. We appended this in the revised text (Line 172).

In the next part of the manuscript, the authors assessed the function of this Nef-G63E mutant. The rational to introduce Ferritin in this part of the document is not clear for the reader. Furthermore, a subgroup for each (NAb+ versus NAb-) is shown: 4 for NAbneg versus 6 for NAbpos.

We appreciate the point. As introduced, Swingler et al Cell Host Microbe 2008 reported HIV-infected macrophage-derived ferritin as a potentially B cell-disrupting factor. In that paper, viral load, ferritin and binding antibody titers positively correlated. Current data shows that SIVmac239-specific NAb induction is distinct from such kinetics already versus viral load (Fig 3-Supplement 1C), and ferritin levels were measured for some available samples more simply for confirmation. We appended three more available samples in the NAb- group. (The six NAb+/G63E animals correspond to the ones with B-cell data in Figure 7.) Statistical results appear unaffected and robust, as shown in this version. The revised manuscript incorporates appended explanation for the former.

Similarly, whereas the authors observed a role of nef mutant on pAkt Ser473 (less induced) in comparison to WT, the authors suggest that this may have an impact on T cell survival.

We appreciate the point. In the first submission we obtained peripheral memory Tfh decrease, whereas it is true that this is indirect. In the current revision we have addressed apoptotic cell death, shown to increase with Nef-G63E mutation (Figure 4F).

The rational to analyze CXCR3-CXCR5+PD-1+ memory follicular Th (Tfh) is not clear. Moreover, the references used are not the adequately cited. Indeed, these papers show an expansion. See the literature for a depletion (Xu H, J Immunol. 2015; Moukambi F, PLoS Pathog. 2015; Yamamoto T, Sci Transl Med. 2015; Xu H, J Immunol. 2018 Moukambi F, Mucosal Immunol. 2019).

We appreciate these points on in vivo CD4+ T cells.

Peripheral memory Tfh was reported to correlate with Ab cross-reactivity in one human cohort (Locci et al, Immunity 2013) and we concisely examined the subset in the current NAb induction. We mentioned this in the revised manuscript.

Moukambi F et al, PLoS Pathog 2015 & Mucosal Immunol 2019 are demonstrative work on acute-phase destruction. We have cited non-neonatal/vaccine-related ones suggested, including these two, in the revised manuscript. The biphasic dysregulation of Th (acute-phase destruction and chronic-phase adverse hyper-expansion) may indeed have a unique role with the current phenotype, which is beyond aim of the current analysis. We have concisely mentioned on this in the Discussion.

Then, the authors assess the potential B-cell-intrinsic influence of the G63E-Nef phenotype. The rational here is clearly indicated, making sense with figure 1. Furthermore, this part is clearer. The dot-plots merit to be revised and the markers used better stated. The authors indicate that Nef invasion upregulates pAkt Ser473 assuming aberrant PI3K/mTORC2 signaling. What is the impact of Nef-G63E mutant on pAkt Ser473 using in vitro model of transfer. This is not addressed for comparison.

We appreciate the remarks/suggestions, also pointed out by Reviewers 1 and 2. We have performed three sets of in vitro reconstitution experiments visually and functionally addressing how Nef transfer to B cells can be modulated (new Fig 6), and edited text accordingly.

Minor points are:

- the presence of references in the legend.

-some Ab clones are in the table, however they are not used such CD38 and CD138, which are well known to be non-valid B cell markers for monkeys."

We appreciate the suggestions.

Mentioning on reference have been removed from the legend (Fig.1, Fig. 3) and moved to the corresponding Methods section (Fig. 1).

We also understood this well in advance (CD38/CD138), and incorporated them in the memory B-cell panel just to check whether they ever behave in a specific pattern. As expected, no notable behavior was observed in these NAb inducers.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: HIV and simian immunodeficiency virus (SIV) infections are known for impaired neutralizing antibody (NAb) responses. While sequential virus-host B cell interaction appears to be basally required for NAb induction, driver molecular signatures predisposing to NAb induction still remain largely unknown. Here we describe SIV-specific NAb induction following a virus-host interplay reciprocal to a congenital human antibody-dysregulating disease. Screening of seventy neutralization-resistant SIVmac239-infected macaques found nine NAb-inducing animals, with seven selecting for a specific CD8+ T-cell escape mutation in viral nef before NAb induction. This mutation reduced aberrant Nef interaction-mediated drive of B-cell maturation-limiting phosphoinositide 3-kinase (PI3K)/mammalian target of rapamycin complex 2 (mTORC2). In vivo imaging cytometry depicted preferential Nef perturbation of cognate Envelope-specific B cells, corroborating cognate B-cell maturation post-mutant selection up to NAb induction. Results exemplify a NAb induction pattern extrinsically reciprocal to human PI3K gain-of-function antibody dysregulation, shaped through two-front, sequential virus interaction with both wings of adaptive immunity.

Revised version: HIV and simian immunodeficiency virus (SIV) infections are known for impaired neutralizing antibody (NAb) responses. While sequential virus-host B cell interaction appears to be basally required for NAb induction, driver molecular signatures predisposing to NAb induction still remain largely unknown. Here we describe SIV-specific NAb induction following a virus-host interplay decreasing aberrant viral drive of phosphoinositide 3-kinase (PI3K). Screening of seventy difficult-to-neutralize SIVmac239-infected macaques found nine NAb-inducing animals, with seven selecting for a specific CD8+ T-cell escape mutation in viral nef before NAb induction. This Nef-G63E mutation reduced excess Nef interaction-mediated drive of B-cell maturation-limiting PI3K/mammalian target of rapamycin complex 2 (mTORC2). In vivo imaging cytometry depicted preferential Nef perturbation of cognate Envelope-specific B cells, suggestive of polarized contact-dependent Nef transfer and corroborating cognate B-cell maturation post-mutant selection up to NAb induction. Results collectively exemplify a NAb induction pattern extrinsically reciprocal to human PI3K gain-of-function antibody-dysregulating disease, and indicate that harnessing the PI3K/mTORC2 axis may facilitate NAb induction against difficult-to-neutralize viruses including HIV/SIV.

Introduction

Earlier version: Virus-specific neutralizing antibody (NAb) responses by B cells are induced by a tight cooperation of adaptive immune cells (Kumar et al., 2010; Shulman et al., 2014; Giltin et al., 2014; Wang et al., 2014) and often play a central role in clearance of acute viral infections (Junt et al., 2007). In contrast, persistence-prone viruses such as human immunodeficiency virus type 1 (HIV-1), simian immunodeficiency virus (SIV) and lymphocytic choriomeningitis virus (LCMV) variously equip themselves with B cell/antibody-inhibitory countermeasures (Moir et al., 2001; Mattapallil et al., 2005; Sommerstein et al., 2015; Sammicheli et al., 2016; Fallet et al., 2016; Mason et al., 2016) and NAb responses are impaired (Levesque et al.,2009; Mikell et al., 2011). These viruses successfully suppress elicitation of potent NAb responses particularly in acute infection (Hunziker et al., 2003; Tomaras et al., 2008) and establish viral persistence, posing considerable challenge for developing protective strategies.

Revised version: Virus-specific neutralizing antibody (NAb) responses by B cells are induced by a tight cooperation of adaptive immune cells (Kumar et al., 2010; Shulman et al., 2014; Giltin et al., 2014; Wang et al., 2014) and often play a central role in clearance of acute viral infections (Junt et al., 2007). In contrast, persistence-prone viruses such as human immunodeficiency virus type 1 (HIV-1), simian immunodeficiency virus (SIV) and lymphocytic choriomeningitis virus (LCMV) variously equip themselves with B cell/antibody-inhibitory countermeasures (Moir et al., 2001; Mattapallil et al., 2005; Sommerstein et al., 2015; Sammicheli et al., 2016; Fallet et al., 2016; Mason et al., 2016) and NAb responses are impaired (Levesque et al., 2009; Mikell et al., 2011). These viruses successfully suppress elicitation of potent NAb responses particularly in acute infection (Hunziker et al., 2003; Tomaras et al., 2008) and establish viral persistence, posing considerable challenge for developing protective strategies. Particularly, HIV and SIV establish early a large in vivo body of infection with a distinct host genome-integrating retroviral life cycle (Whitney et al., 2014). In addition, lentiviruses HIV/SIV are unique in launching a matrix of host immune-perturbing interactions mainly by their six remarkably pleiotropic accessory viral proteins optimally fueling pathogenesis (Kirchhoff et al., 1995; Sheehy et al., 2002; Harris et al., 2003; Schindler et al., 2006; Neil et al., 2008; Zhang et al., 2009; Laguette et al., 2011; Yamada et al., 2018; Joas et al., 2018; Langer et al., 2019; Yan et al., 2019; Joas et al., 2020; Khan et al., 2020; Volcic et al., 2020; Reuschl et al., 2022). A body of evidence has depicted this over these 30 years, whereas its entity, including focal mechanisms of humoral immune perturbation in HIV/SIV infection, still remains elusive.

Earlier version: Reciprocal to this detrimental acute-phase NAb absence, we and others have previously described that acute-phase passive NAb infusion can trigger an endogenous T-cell synergism resulting in robust SIV/SHIV (simian/human immunodeficiency virus) replication control (Haigwood et al., 1996; Yamamoto et al., 2007; Ng et al., 2010; Iseda et al., 2016; Nishimura et al., 2017). This shows that virus-specific NAbs not only confer sterile protection but also can evoke T-cell-mediated non-sterile viral control, suggesting the importance of inducing endogenous NAb responses during an optimal time frame as well as humoral-cellular response synergisms. Therefore, identifying mechanisms driving NAb induction against such viruses is an important step to eventually design NAb-based HIV control strategies.

Revised version: Contrasting this detrimental consequence of acute-phase endogenous NAb absence driven by adverse virus-host interactions in early HIV/SIV infection, we and others have previously described in in vivo experimental models that acute-phase passive NAb infusion can trigger an endogenous T-cell synergism resulting in robust SIV/SHIV (simian/human immunodeficiency virus) replication control (Haigwood et al., 1996; Yamamoto et al., 2007; Ng et al., 2010; Iseda et al., 2016; Nishimura et al., 2017). This shows that virus-specific Nabs not only confer sterile protection but also can evoke T-cell-mediated non-sterile viral control, suggesting the importance of inducing endogenous NAb responses during an optimal time frame as well as humoral-cellular response synergisms. Therefore, identifying mechanisms driving NAb induction against such viruses is an important step to eventually design NAb-based HIV control strategies.

Earlier version: As an approach, analysis of in vivo immunological events linked with NAb induction against difficult-to-neutralize SIVs in a non-human primate model can provide important insights into this goal. HIV-specific NAb induction patterns and correlates, such as antibody-NAb coevolution (Moore et al., 2012), autoimmune-driven induction (Moody et al., 2016) and natural killer cell-related host polymorphisms (Bradley et al., 2018) have been variously reported to date; the wealth of knowledge on this topic collectively, and interestingly, rather suggests that pathways to NAb induction against neutralization-resistant viruses including HIV/SIV can be highly redundant and there may be as-yet-unknown in vivo mechanisms that can robustly drive NAb induction. For example, the neutralization resistance of certain SIV strains apparently is not explanatory by the aforementioned, posing such models as attractive analytic targets of NAb induction mechanisms.

Revised version: As an approach, analysis of in vivo immunological events linked with NAb induction against difficult-to-neutralize SIVs in a non-human primate model can provide important insights into this goal. Various in vivo signatures of HIV-specific NAb induction, such as antibody-NAb coevolution (Moore et al., 2012), autoimmune-driven induction (Moody et al., 2016) and natural killer cell-related host polymorphisms (Bradley et al., 2018) have been reported to date. The variety collectively, and interestingly, indicates that pathways to NAb induction against difficult-to-neutralize viruses including HIV/SIV are redundant, conversely implicating as-yet-unknown mechanisms driving NAb induction. For example, the neutralization resistance of certain SIV strains apparently is not explanatory by any of the aforementioned, posing such models as attractive analytic targets of NAb induction mechanisms.

Earlier version: In the present study, we examined virus-specific antibody responses in rhesus macaques infected with a highly neutralization-resistant SIV strain, SIVmac239. This virus is highly pathogenic, shows persistent viremia and hosts lack NAb responses throughout infection (Kestler et al., 1991). Yet, upon performing a large-scale screening of SIVmac239-infected rhesus macaques for up to 100 weeks, we noticed a subgroup inducing NAbs in the chronic phase. Interestingly, these animals commonly selected for a specific CD8+ cytotoxic T lymphocyte (CTL) escape mutation in the viral Nef-coding gene before NAb induction. The mutant Nef, compared with wild-type (WT) Nef, manifested a decrease in aberrant interaction with phosphoinositide 3-kinase (PI3K)/mammalian target of rapamycin complex 2 (mTORC2), resulting in decreased downstream hyperactivation of the canonical B-cell negative regulator Akt (Omori et al., 2006; Limon et al., 2014). Nef was graphically identified via machine learning-comprising imaging cytometry to preferentially target Env-specific B cells in vivo, and in those macaques inducing NAbs, the NAb induction was linked with sustained Env-specific B-cell responses after the mutant Nef selection. Results describe a B-cell functional boosting reciprocal to human PI3K gain-of-function antibody dysregulation, and suggest that intervening PI3K/mTORC2 signaling can potentially result in harnessing NAb induction against difficult-to-neutralize viruses.

Revised version: In the present study, we examined virus-specific antibody responses in rhesus macaques infected with a highly difficult-to-neutralize SIV strain, SIVmac239. This virus is pathogenic in rhesus macaques causing simian AIDS across a broad range of geographical origin of macaques (Cumont et al., 2008), shows persistent viremia and hosts lack NAb responses throughout infection (Kestler et al., 1991; Nomura et al., 2012). Yet, upon performing a large-scale screening of SIVmac239-infected Burmese rhesus macaques for up to 100 weeks, we noticed a subgroup inducing NAbs in the chronic phase. Interestingly, these animals commonly selected for a specific CD8+ cytotoxic T lymphocyte (CTL) escape mutation in the viral Nef-coding gene before NAb induction. The mutant Nef, compared with wild-type (WT) Nef, manifested a decrease in aberrant interaction with phosphoinositide 3-kinase (PI3K)/mammalian target of rapamycin complex 2 (mTORC2), resulting in decreased downstream hyperactivation of the canonical B-cell negative regulator Akt (Omori et al., 2006; Limon et al., 2014). Nef was graphically identified via machine learning-assisted imaging cytometry to preferentially target Env-specific B cells in vivo, and in those macaques inducing NAbs, the NAb induction was linked with sustained Env-specific B-cell responses after or during the mutant Nef selection. Results describe a B-cell functional boosting phenotypically reciprocal to a recently found human PI3K gain-of-function and antibody-dysregulating inborn error of immunity (IEI), activated PI3 kinase delta syndrome (APDS) (Angulo et al., 2013; Lucas et al., 2014), and suggest that intervening PI3K/mTORC2 signaling can potentially result in harnessing NAb induction against difficult-to-neutralize viruses.

Results

Identification of macaques inducing SIVmac239-neutralizing antibodies

Earlier version: We performed a retrospective antibody profile screening in rhesus macaques infected with NAb-resistant SIVmac239 (n = 70) (Figure 1A) and identified a group of animals inducing anti-SIVmac239 NAb responses (n = 9), which were subjected to characterization. These NAb inducers showed persistent viremia with no significant difference in viral loads compared with a profile-clarified naïve, major histocompatibility complex class I (MHC-I) haplotype-balanced subgroup of NAb non-inducers (n = 19) representatively utilized for comparison hereafter (Figure 1B and Figure 1-figure supplement 1). Plasma SIVmac239-NAb titers measured by 10 TCID50 SIVmac239 virus-killing assay showed an average maximum titer of 1:16, being induced at an average of 48 weeks post-infection (p.i.) (Figure 1C). Two of the nine NAb inducers showed detectable NAb responses before six months p.i., while the remaining seven induced NAbs after six months. Anti-SIVmac239 neutralizing activity was confirmed in immunoglobulin G (IgG) purified from plasma of these NAb inducers (Figure 1-figure supplement 2). SIV Env-binding IgGs developed from early infection in NAb inducers as well as NAb non-inducers (Figure 1-figure supplement 3); titers were not higher but rather lower at year 1 p.i. in the NAb inducers (Figure 1-figure supplement 3B and 3C), consistent with other reports (Havenar-Daughton et al., 2016). NAb inducers and non-inducers showed similar patterns of variations in viral env sequences (Burns et al., 1993), mainly in variable regions 1, 2 and 4 (V1, V2 and V4) (Figure 1-figure supplement 4).

Revised version: We performed a retrospective antibody profile screening in rhesus macaques infected with NAb-resistant SIVmac239 (n = 70) (Figure 1A) and identified a group of animals inducing anti-SIVmac239 NAb responses (n = 9), which were subjected to characterization. These NAb inducers showed persistent viremia with no significant difference in viral loads compared with a profile-clarified naïve, major histocompatibility complex class I (MHC-I) haplotype-balanced subgroup of NAb non-inducers (n = 19) representatively utilized for comparison hereafter (Figure 1B and Figure 1-figure supplement 1). Plasma SIVmac239-NAb titers measured by 10 TCID50 SIVmac239 virus-killing assay showed an average maximum titer of 1:16, being induced at an average of 48 weeks post-infection (p.i.) (Figure 1C). Two of the nine NAb inducers showed detectable NAb responses before six months p.i., while the remaining seven induced NAbs after six months. Anti-SIVmac239 neutralizing activity was confirmed in immunoglobulin G (IgG) purified from plasma of these NAb inducers (Figure 1-figure supplement 2). SIV Env-binding IgGs developed from early infection in NAb inducers as well as non-rapid-progressing NAb non-inducers, the latter differing from rapid progressors known to manifest serological failure (Hirsch et al., 2004; Nakane et al., 2013) (Figure 1-figure supplement 3A); titers were not higher but rather lower at year 1 p.i. in the NAb inducers (Figure 1-figure supplement 3B and 3C). This was consistent with other reports (Havenar-Daughton et al., 2016) and differing with gross B-cell enhancement in Nef-deleted SIV infection with enhanced anti-SIV binding antibody titers accompanying marginal neutralization (Adnan et al., 2016). NAb inducers and non-inducers showed similar patterns of variations in viral env sequences (Burns et al., 1993), mainly in variable regions 1, 2 and 4 (V1, V2 and V4) (Figure 1-figure supplement 4).

NAb induction against NAb-resistant SIVmac239

(A) Study design.

Earlier version: (B) Plasma viral loads (SIV gag RNA copies/ml plasma) in NAb non-inducers (left) and inducers (right). Viral loads have been previously partially reported (Yamamoto et al., 2007; Nomura et al., 2012; Takahashi et al., 2013; Nakane et al., 2013; Iwamoto et al., 2014).

Revised version: (B) Plasma viral loads (SIV gag RNA copies/ml plasma) in NAb non-inducers (left) and inducers (right).

Earlier version: (C) Plasma SIVmac239 100% neutralizing end point titers by 10 TCID50 killing assay on MT4-R5 cells. Points on the dotted line show marginally NAb-positive results (< 1:2). In some animals, titers were comparable with our reported results using MT4 cells (Kawada et al., 2007).

Revised version: (C) Plasma SIVmac239 100% neutralizing end point titers by 10 TCID50 killing assay on MT4-R5 cells. Points on the dotted line show marginally NAb-positive results (< 1:2). In some animals, titers were comparable with our reported results using MT4 cells.

Selection of a viral nef mutation (Nef-G63E) precedes chronic-phase SIVmac239-specific NAb induction

Earlier version: Viral nonsynonymous polymorphisms outside env were next assessed. Strikingly, we found selection of a viral genome mutation resulting in G (glycine)-to-E (glutamic acid) substitution at residue 63 of Nef (Nef-G63E) in seven of the nine NAb inducers (Figure 2A). This 63rd residue is in the unstructured N-terminus of Nef, flanked by conserved α-helices (Figure 2B). This region is conserved in HIV-2/SIV but not in HIV-1, occasionally being deleted in laboratory HIV-1 strains (Geyer et al., 2001). AlphaFold2-based structure prediction did not derive palpable change except for low-probability disruption of the alpha helices (data not shown). The Nef G63E mutation was not observed in two early NAb inducers (with detectable NAbs before 6 months p.i.) but selected in all of the remaining seven inducing detectable NAb responses after 6 months. In contrast, this mutation was found only in two of the nineteen NAb non-inducers, including one rapid progressor (Figure 1-figure supplement 1) and selection was significantly enriched in the NAb inducers (Figure 2C). Analysis of Nef-G63E mutation frequencies in plasma viruses showed that Nef-G63E selection preceded or at least paralleled NAb induction (Figure 2D).

Revised version: Viral nonsynonymous polymorphisms outside env were next assessed. Strikingly, we found selection of a viral genome mutation resulting in G (glycine)-to-E (glutamic acid) substitution at residue 63 of Nef (Nef-G63E) in seven of the nine NAb inducers (Figure 2A). This 63rd residue lies in the unstructured N-terminus of Nef, flanked by two α-helices conserved in SIVmac/SIVsmm (Figure 2B). This region is generally polymorphic among HIV-1/HIV-2/SIV, occasionally being deleted in laboratory and primary isolate HIV-1 strains (Geyer et al., 2001). AlphaFold2-based structure prediction did not derive palpable change except for low-probability disruption of the alpha helices (data not shown). The Nef G63E mutation was not observed in two early NAb inducers (with detectable NAbs before 6 months p.i.) but selected in all of the remaining seven inducing detectable NAb responses after 6 months. In contrast, this mutation was found only in two of the nineteen NAb non-inducers, including one rapid progressor (Figure 1-figure supplement 1) and selection was significantly enriched in the NAb inducers as compared to the nineteen non-inducers (Figure 2C). Analysis of Nef-G63E mutation frequencies in plasma viruses showed that Nef-G63E selection preceded or at least paralleled NAb induction (Figure 2D).

Selection of a viral nef mutation (Nef-G63E) before NAb induction.

(A) Viral nef mutations in NAb inducers. A linear schema indicating Nef functional domains is aligned above. Mutations at timepoints with NAb- (indicated by white wedge; mostly before 6 months p.i.), NAb+/- (green wedge) and NAb+) (purple wedge; mostly after 1 year) are shown in individual animals. Black and dark gray represent dominant and subdominant mutations (or deletion in R-360) by direct sequencing, respectively. Red and pink indicate dominant and subdominant G63E detection, respectively.

(B) Schema of SIVmac239 Nef structure and Nef-G63E mutation orientation.

(C) Comparison of frequencies of macaques having Nef-G63E in plasma viruses between NAb non-inducers and inducers. Compared by Fisher’s exact test.

Earlier version: (D) Temporal relationship of Nef G63E frequencies in plasma virus and NAb induction. Black boxes (left Y axis) show log2 NAb titers; red triangles and green diamonds (right Y axis) show percentage of G63E and G63R mutations detected by subcloning (15 clones/point on average), respectively.

Revised version: (D) Temporal relationship of Nef G63E frequencies in plasma virus and NAb induction. Black boxes (left Y axis) show log2 NAb titers; red triangles and blue diamonds (right Y axis) show percentage of G63E and G63R mutations detected by subcloning (15 clones/point on average), respectively.

Nef-G63E is a CD8+ T-cell escape mutation

Earlier version: These seven NAb inducers selecting this mutation elicited CD8+ T-cell responses specific for a 9-mer peptide Nef62–70 QW9 (QGQYMNTPW) (Figure 3A). This Nef62–70 QW9 epitope is restricted by MHC-I molecules Mamu-B*004:01 and Mamu-B*039:01 (Evans et al., 1999; Sette et al., 2012). Possession of these accounted for six cases of Nef-G63E selection (Figure 3A), and the remaining one animal possessed Mamu-A1*032:02 predicted to bind to Nef63–70 GW8 (NetMHCpan). Ten of nineteen NAb non-inducers also had either of these alleles (Figure 1-figure supplement 1), indicating that NAb induction is not associated with these MHC-I genotypes (P = 0.25 by Fisher’s exact test, data not shown) but with the Nef-G63E mutation itself (Figure 2C). Nef-G63E was confirmed to be an escape mutation from Nef62-70-specific CD8+ T-cell responses (Figure 3B). NAb non-inducers possessing these alleles elicited little or no Nef62-70-specific CD8+ T-cell responses (Figure 3-figure supplement 1A), suggesting that in vivo selection and fixation of this Nef-G63E SIV was indeed Nef62-70-specific CD8+ T cell-dependent. Replication kinetics of SIV carrying the Nef-G63E mutation was comparable with wild-type (WT) on a cynomolgus macaque HSC-F CD4+ T-cell line (Akari et al., 1996), stating that the mutation does not critically impair viral replication (Figure 3-figure supplement 1B), and plasma viral loads were comparable between Nef-G63E mutant-selecting NAb inducers versus non-inducers (Figure 3-figure supplement 1C). These somewhat contrasted conditional reversion of a more replication-impairing nef mutation related with CD3 downregulation loss in SIVmac239 infection (Joas et al., 2020). These results indicate NAb induction following in vivo selection and fixation of the CD8+ T-cell escape nef mutation, Nef-G63E.

Revised version: These seven NAb inducers selecting this mutation elicited CD8+ T-cell responses specific for a 9-mer peptide Nef62–70 QW9 (QGQYMNTPW) (Figure 3A). This Nef62–70 QW9 epitope is restricted by MHC-I molecules Mamu-B*004:01 and Mamu-B*039:01 (Evans et al., 1999; Sette et al., 2012). Possession of these accounted for six cases of Nef-G63E selection (Figure 3A), and the remaining one animal possessed Mamu-A1*032:02 predicted to bind to Nef63–70 GW8 (NetMHCpan). Ten of nineteen NAb non-inducers also had either of these alleles (Figure 1-figure supplement 1). This did not significantly differ with the NAb inducer group (P = 0.25 by Fisher’s exact test, data not shown), indicating that NAb induction was not simply linked with possession of these MHC-I genotypes but instead required furthermore specific selection of the Nef-G63E mutation (Figure 2C). Nef-G63E was confirmed to be an escape mutation from Nef62-70-specific CD8+ T-cell responses (Figure 3B). NAb non-inducers possessing these alleles elicited little or no Nef62-70-specific CD8+ T-cell responses (Figure 3-figure supplement 1A), suggesting that in vivo selection and fixation of this Nef-G63E SIV was indeed Nef62-70-specific CD8+ T cell-dependent. Replication kinetics of SIV carrying the Nef-G63E mutation was comparable with wild-type (WT) on a cynomolgus macaque HSC-F CD4+ T-cell line (Akari et al., 1996), excluding the possibility that the mutation critically impairs viral replication (Figure 3-figure supplement 1B), and plasma viral loads were comparable between Nef-G63E mutant-selecting NAb inducers versus non-inducers (Figure 3-figure supplement 1C). These results indicate NAb induction following in vivo selection and fixation of the CD8+ T-cell escape nef mutation, Nef-G63E under viral persistence.

Nef-G63E is a CD8+ T-cell escape mutation.

Earlier version: (A) Nef62–70 QW9-specific CD8+ T-cell frequencies and related MHC-I alleles in seven Nab inducers selecting Nef-G63E. Mamu-B*039:01 and Mamu-B*004:01 are known to restrict Nef62–70 QW9 epitope (Evans et al., 1999; Sette et al., 2012) and binding of Nef63–70 peptide to Mamu-A1*032:02 was predicted.

Revised version: (A) Nef62–70 QW9-specific CD8+ T-cell frequencies and related MHC-I alleles in seven NAb inducers selecting Nef-G63E. Mamu-B*039:01 and Mamu-B*004:01 are known to restrict Nef62–70 QW9 epitope and binding of Nef63–70 peptide to Mamu-A1*032:02 was predicted.

Earlier version: (B) CD8+ T-cell responses specific to wild-type Nef62–70 or mutant Nef62–70.G63E or Nef62–70.G63R peptides. See also Figure S3.

Revised version: (B) CD8+ T-cell responses specific to wild-type Nef62–70 or mutant Nef62–70.G63E or Nef62– 70.G63R peptides.

G63E mutation reduces aberrant Nef interaction-mediated drive of PI3K/mTORC2

Earlier version: We next focused on the functional phenotype of this highlighted Nef-G63E mutant SIV in infected cells. An essence of host perturbation by Nef is its wide-spectrum molecular downregulation, ultimately increasing viral replication (Schindler et al., 2004; Zhang et al., 2009). To address possible amelioration, we compared downregulation of major targets CD3, CD4, MHC-I, CXCR4 and BST-2 (Jia et al., 2009) in infected HSC-F cells. Nef-G63E mutation did not confer notable change compared with wild-type (WT) [P = not significant (n.s.) for all molecules, data not shown], implicating other non-canonical changes (Figure 4A). Literature suggested ferritin negative influence on B cells proportionate to HIV-infected macrophage-derived production (Swingler et al., 2008). Here, plasma ferritin showed no differences (Figure 4-figure supplement 1A) and viral loads were comparable between Nef-G63E-selecting NAb inducers and non-inducers (Figure 4-figure supplement 1B and 1C), arguing against gross involvement of ferritin as well as other viral replication-changing Nef phenotypes at least in this model.

Revised version: We next focused on the functional phenotype of this highlighted Nef-G63E mutant SIV in infected cells. An essence of host perturbation by Nef is its wide-spectrum molecular downregulation, ultimately increasing viral replication (Schindler et al., 2004; Zhang et al., 2009). To address possible amelioration, we compared downregulation of major targets CD3, CD4, MHC-I, CXCR4 and BST-2 (Jia et al., 2009) in infected HSC-F cells. Nef-G63E mutation did not confer notable change compared with wild-type (WT) [P = not significant (n.s.) for all molecules, data not shown], implicating other non-canonical changes (Figure 4A). Literature suggested ferritin to exert negative influence on B cells proportionate to HIV-infected macrophage-derived production (Swingler et al., 2008). Here, plasma ferritin levels showed no differences (Figure 4-figure supplement 1A) and viral loads were comparable between Nef-G63E-selecting NAb inducers and non-inducers (Figure 4-figure supplement 1B and 1C), arguing against gross involvement of ferritin as well as other viral replication-changing Nef phenotypes at least in this model.

Nef-G63E mutation reduces PI3K/mTORC2 binding and pAkt drive.

(A) Representative surface expression level histograms of CD3, CD4, MHC class I, CXCR4 and BST-2 in CD4loNef+ subpopulations after wild-type (WT) or Nef-G63E mutant SIV infection at multiplicity of infection (MOI) 0.1 on HSC-F cells.

(B) Left: representative histograms of relative pAkt serine (Ser) 473 levels in p27+ subpopulations after WT or Nef-G63E mutant SIV infection at MOI 0.1 on HSC-F cells. Numbers show pAkt Ser473 mean fluorescence intensities (MFIs) for each. Right: Deviation of pAkt Ser473 MFIs in p27+ HSC-F cells compared with mean MFI of uninfected cells. Compared by unpaired t test.

(C) Left: Relative pAkt Ser473 levels (normalized to mean MFI of uninfected controls) in Nef+ HSC-F cells assessed for serum starvation (MOI 0.2, 1 day p.i.). Adjusted P values show results of comparison via Sidak’s post-hoc test of 2-way ANOVA (C, left). Right: Deviation of pAkt Ser473 MFIs in Nef+ HSC-F cells assessed for high-MOI infection (MOI 5, 1-day p.i.). Compared by unpaired t test.

Earlier version: (D) Proximity ligation assay (PLA) of Nef binding with mTOR, GβL/mLST8, PI3Kp85 and PI3Kp110α. MFI-based binding index was calculated as (anti-Nef/anti-partner) - (isotype/anti-partner) - (anti-Nef/isotype) + (isotype/isotype). Histograms for samples and isotype/isotype are representatively shown. Differences in MFI binding indexes can be enhanced compared with comparison of raw MFIs. Compared by paired t tests. Data represent one of two [(A); (B); (C, right)] independent experiments in quadruplicate, four independent experiments performed in triplicate [(C, left)] or four independent single-well comparison experiments pooled for statistical analysis (D). Bars: mean ± SD (B, C right), mean ± SEM (C left).

Revised version: (D) Proximity ligation assay (PLA) of Nef binding with mTOR, GβL/mLST8, PI3Kp85 and PI3Kp110α. MFI-based binding index was calculated as (anti-Nef/anti-partner) - (isotype/anti-partner) - (anti-Nef/isotype) + (isotype/isotype). Histograms for samples and isotype/isotype are representatively shown. Differences in MFI binding indexes can be enhanced compared with comparison of raw MFIs. Compared by paired t tests.

Added in the revised version: (E) Sin1 co-immunoprecipitation analysis of WT versus Nef-G63E mutant SIV-infected HSC-F cells (infected at MOI 0.05, day 3). Immunoblotting of Nef (37 kDa) in whole cell lysates (lanes 4-6) and anti-Sin1 antibody immunoprecipitates (lanes 1-3) are shown.

Added in the revised version: (F) Cell death frequencies of SIV-infected cells measured by Annexin V positivity (% of CD4lo) (infected at MOI 0.1, day 3). Compared by unpaired t test.

Added in the revised version: Data represent one of two [(A); (B); (C, right)] independent experiments in quadruplicate, four independent experiments performed in triplicate [(C, left)], four independent single-well comparison experiments pooled for statistical analysis (D), two experiments (E) or two experiments performed with six wells/control (F). Bars: mean ± SD (B, C right), mean ± SEM (C left, F).

Earlier version: Thus, we took interest in immunosignaling-modulating properties of Nef-G63E SIV. Analysis of Akt, the predominant immune-intrinsic negative brake of B-cell maturation and AFC/antibody responses in vivo (Omori et al., 2006; Limon et al., 2014; Fruman et al., 1999; Ray et al., 2015; Sander et al., 2015), revealed that its serine 473 phosphorylation, non-canonically known for Nef-mediated upregulation (Kumar et al., 2016), was significantly lower on day 3 after Nef-G63E SIV infection at low multiplicity of infection (MOI) compared with WT (Figure 4B), stronger than histogram deviation levels in PI3K gain-of-function mice with full disease phenotype (Avery et al., 2018). The threonine 308-phosphorylated Akt (pAkt Thr308) level was more unaffected (Figure 4-figure supplement 1B).

Revised version: Thus, we took interest in immunosignaling-modulating properties of Nef-G63E SIV. Analysis of Akt, the predominant immune-intrinsic negative brake of B-cell maturation and AFC/antibody responses in vivo (Omori et al., 2006; Limon et al., 2014; Fruman et al., 1999; Ray et al., 2015; Sander et al., 2015), revealed that its serine 473 phosphorylation, non-canonically known for Nef-mediated upregulation (Kumar et al., 2016), was significantly lower on day 3 after Nef-G63E SIV infection at low multiplicity of infection (MOI) compared with WT (Figure 4B), stronger than histogram deviation levels in PI3K gain-of-function mice with full recapitulation of B-cell/antibody-dysregulating human APDS disease phenotype (Avery et al., 2018). The threonine 308-phosphorylated Akt (pAkt Thr308) level remained unaffected (Figure 4-figure supplement 1B).

Earlier version: We further assessed properties of serine 473-phoyphorylated Akt (pAkt Ser473) inhibition by Nef-G63E SIV, and found that external PI3K stimuli (serum starvation) (Kennedy et al., 1997) accelerated phenotype appearance from day 3 to day 1 post-infection (Figure 4C, left). A similar trend was obtained by short-term PI3K stimulation with interferon-γ (Nguyen et al., 2001), interleukin-2 (Marzec et al., 2008) and SIV Env (François & Klotman, 2003) (Figure 4-figure supplement 1C). Env-stimulating high-MOI infection also accelerated phenotype appearance, with enhanced 50% reduction (Figure 4C, right). These showed that the Akt-inhibitory G63E mutant Nef phenotype is PI3K stimuli-dependent. Transcriptome analysis signatured decrease in PI3K-pAkt-FoxO1-related genes in mutant SIV infection (Figure 4-figure supplement 1D). As Akt is also a survival mediator this implicates that Nef-G63E mutation may be Th-cytopathic, decreasing Th dysfunctional expansion in LCMV (Recher et al., 2004) and HIV/SIV (Gray et al., 2011; Lindqvist et al., 2012; Petrovas et al., 2012) infections. Here, reduced peripheral CXCR3-CXCR5+PD-1+ memory follicular Th (Tfh) (Locci et al., 2013) was in keeping (Figure 4-figure supplement 1E and 1F).

Revised version: We further assessed properties of serine 473-phoyphorylated Akt (pAkt Ser473) inhibition by Nef-G63E SIV, and found that external PI3K stimuli (serum starvation) (Kennedy et al., 1997) accelerated phenotype appearance from day 3 to day 1 post-infection (Figure 4C, left). A similar trend was obtained by short-term PI3K stimulation with interferon-γ (Nguyen et al., 2001), interleukin-2 (Marzec et al., 2008) and SIV Env (François & Klotman, 2003) (Figure 4-figure supplement 1C). A high-MOI SIV infection, comprising higher initial concentration of extracellular Env stimuli, also accelerated phenotype appearance from day 3 to day 1 post-infection, with stronger pAkt reduction (Figure 4C, right). These showed that the Akt-inhibitory G63E mutant Nef phenotype is PI3K stimuli-dependent. Transcriptome analysis signatured decrease in PI3K-pAkt-FoxO1-related genes in mutant SIV infection (Figure 4-figure supplement 1D). As Akt is also a survival mediator this implicated that Nef-G63E mutation may be Th-cytopathic, decreasing Th dysfunctional expansion (Baumjohann et al., 2013) as had been observed in LCMV strain WE (Recher et al., 2004) and chronic HIV/SIV (Gray et al., 2011; Lindqvist et al., 2012; Petrovas et al., 2012) infections. Here, reduced chronic-phase peripheral CXCR3-CXCR5+PD-1+ memory follicular Th (Tfh), linked with antibody cross-reactivity in one cohort (Locci et al., 2013), was somewhat in line with this notion (Figure 4-figure supplement 1E and 1F). We additionally examined another CTL escape mutant Nef-G63R, selected with marginal statistical significance primarily in early stage (Figure 4-figure supplement 2A). This mutant did not associate with NAb induction in one control animal (R06-034) even upon preferential selection (Figure 4-figure supplement 2B) and Nef-G63R mutant was not similarly decreased in pAkt drive (Figure 4-figure supplement 2C), implicating that the Nef-G63E mutation was more tightly linked with Nab induction in terms of an immunosignaling phenotype.

Earlier version: We further explored molecular traits of this decreased Akt hyperactivation. We reasoned that comparing endogenous Nef binding patterns would be adequate and analyzed SIV-infected HSC-F cells with a flow cytometry-based proximity ligation assay (PLA) (Leuchowius et al., 2009; Avin et al., 2017), by calculating a mean fluorescence intensity (MFI) deviation-based binding index for each. We found that this Nef-G63E mutation causes significant decrease in Nef binding to PI3K p85/p110α and downstream mTORC2 components mTOR (Sarbassov et al., 2005) and GβL/mLST8 (Kim et al., 2003) in the CD4lo-SIV Gag p27+ infected population (Figure 4D). Collectively, the Nef-G63E mutation results in decreased aberrant Nef bridging and signal transduction of PI3K/mTORC2, explaining decreased substrate Akt Ser473 phosphorylation by mTORC2. Thus, Nef-G63E SIV is a mutant virus decreased in aberrant interaction/drive of B-cell-inhibitory PI3K/mTORC2 signaling. This molecular signature was reciprocal to a recently found human inborn error of immunity (IEI), activated PI3 kinase delta syndrome (APDS) (Angulo et al., 2013; Lucas et al., 2014; Avery et al., 2018).

Revised version: We further explored molecular traits of this decreased Akt hyperactivation. We reasoned that comparing endogenous Nef binding patterns would be adequate and analyzed SIV-infected HSC-F cells with a flow cytometry-based proximity ligation assay (PLA) (Leuchowius et al., 2009; Avin et al., 2017), by calculating a mean fluorescence intensity (MFI) deviation-based binding index for each. We found that this Nef-G63E mutation causes significant decrease in Nef binding to PI3K p85/p110α and downstream mTORC2 components mTOR (Sarbassov et al., 2005) and GβL/mLST8 (Kim et al., 2003) in the CD4lo-SIV Gag p27+ infected population (Figure 4D). Results were biochemically confirmed by a decrease in G63E-mutated Nef binding to the mTORC2-intrinsic cofactor Sin1 in coimmunoprecipitation analysis of infected HSC-F cells (Figure 4E). Collectively, the Nef-G63E mutation results in decreased aberrant Nef bridging and signal transduction of PI3K/mTORC2, explaining decreased substrate Akt Ser473 phosphorylation by mTORC2. These properties were in keeping with an acceleration in CD4-downregulated/infected HSC-F cell apoptosis upon Nef-G63E SIV infection (Figure 4F). Thus, Nef-G63E SIV is a mutant virus decreased in aberrant interaction/drive of B-cell-inhibitory PI3K/mTORC2 signaling, manifesting a molecular signature reciprocal to human APDS (Angulo et al., 2013; Lucas et al., 2014; Avery et al., 2018).

Targeting of lymph node Env-specific B cells by Nef in vivo

Earlier version: Next, we analyzed in vivo targeting of virus-specific B cells by Nef in lymph nodes to explore the potential B-cell-intrinsic influence of the G63E-Nef phenotype. Previously suggested influence of soluble Nef itself (Qiao et al., 2006) and/or related host factors (Swingler et al., 2008) may derive generalized influence on B cells; however, SIV antigen-specific binding antibody responses were rather decreased in NAb inducers (Figure 1-figure supplement 1), viremia was comparable (Figure 3-figure supplement 1C) and ferritin levels were not altered (Figure 4-figure supplement 1A), arguing against these. Therefore, we surmised that some targeted Nef intrusion against Env-specific B cells may be occurring and that the decreased aberrant Nef-PI3K/mTORC2 drive may result in their enhanced lymph node maturation.

Revised version: Next, we analyzed in vivo targeting of virus-specific B cells by Nef in lymph nodes to explore the potential B-cell-intrinsic influence of the G63E-Nef phenotype. Previously suggested influence of soluble Nef itself (Qiao et al., 2006) and/or related host factors (Swingler et al., 2008) may derive generalized influence on B cells; however, SIV antigen-specific binding antibody responses were rather decreased in NAb inducers (Figure 1-figure supplement 1), viremia was comparable (Figure 3-figure supplement 1C) and ferritin levels were not altered (Figure 4-figure supplement 1A), differing from the aforementioned literature stating positive correlation between the three. Therefore, we surmised that some targeted Nef intrusion against Env-specific B cells may be occurring and that the decreased aberrant Nef-PI3K/mTORC2 drive may result in their enhanced lymph node maturation.

While reports (Xu et al., 2009) histologically proposed Nef B-cell transfer, quantitative traits, e.g., invasion frequency and influence on virus-specific B cells, have remained unvisited. Intracellular Nef staining is dim (particularly for B cell-acquired Nef), difficult to examine by conventional flow cytometry. Defining staining cutline requires single-cell images, overcoming confounding of high Nef false-staining via pre-permeation rupture and post-permeation confounding as judged by biologically discontinuous staining and/or batch-inflated signals. Collectively, at the expense of spatial information, we judged that sophisticating imaging cytometry would best visualize Nef-mediated B-cell perturbation in vivo. We analyzed lymph node B cells with Image Stream X MKII, with high-power (> 10-fold) antigen detection [e.g., molecules of equivalent soluble fluorochrome (MESF) 5 vs. MESF 80 in an average flow cytometer for FITC detection] ideal for detection/single-cell verification.

We designed a triple noise cancellation strategy to overcome the above. Firstly, amine reactive dye staining (Perfetto et al., 2006) gated out B cells being Nef-positive due to pre-experimental membrane damage (Figure 5-figure supplement 1A, left). Next, we deployed secondary quantitative parameters of Nef signals deriving from each pixel of the images. Definition of a gray-level cooccurrence matrix (GLCM) (Haralick et al., 1973) numerating adjacent signal deviation as their frequencies on a transpose matrix enables calculation of a variety of feature values summarizing traits of the whole image. A biological assumption of Nef signal continuity in a true staining suggested that the sum of weighted square variance of GLCM, i.e., the {sum of [square of (value – average signal strength) × frequency of each value occurrence]} would separate natural versus artificial Nef signals. This value, Haralick variance, was computed for multiple directions and averaged, deriving Haralick variance mean (Figure 5-figure supplement 1A, X axis), proportionate with staining image unnaturalness. Finally, Nef signal intensity threshold (Figure 5-figure supplement 1A, Y axis) gated out overtly-stained cells void for true-false verification. These excluded B cells showing non-specific, strong-signal binary-clustered staining pixels for Nef deriving a large summated variance (X axis), and/or cells batch-stained for Nef (Y axis), likely originating from post-experimental membrane damage.

Added in the revised version: With this approach we acquired images of a Live/Dead--Nef signal Haralick variance meanlo-Nef signal intensity thresholdlo-Nef+CD19+ B-cells, with fine-textured gradation of Nefint- lo staining with continuity from membrane-proximal regions without sporadic staining speckles (Figure 5A and Figure 5-figure supplement 1). A linear discriminant analysis-based machine learning module, scoring segregation of typical void/valid images, showed this gating to provide the highest two-dimensional separation (Figure 5-figure supplement 1). Combined with biotinylated Env/streptavidin staining, this reproducibly obtained images of Env-specific (Ch 12) CD19+ (Ch 11) B cells without membrane ruptures (Ch 08) and showing fine-textured transferred Nef (Ch 02) (Figure 5B). Nef invasion upregulated pAkt Ser473 (Ch 03) to a range resembling in vitro analysis (Figure 5C), stating that Nef-driven aberrant PI3K/mTORC2 signaling does occur in Nef-invaded B cells in vivo and the current phenotypic change likely affects B-cell dynamics. Strikingly, lymph node Env-specific B cells showed significantly higher Nef-positive frequencies as compared with batch non-Env-specific B cells (Figure 5D). This indicated that infected cell-generated Nef preferentially targets adjacent Env-specific B cells, putatively through infected CD4+ cell-to-B-cell contact/transfer (Xu et al., 2009; Hashimoto et al., 2016) and that a Nef phenotypic change likely dictates Env-specific B-cell maturation.

Preferential targeting of lymph node Env-specific B cells by Nef in vivo.

(A) Representative gating of triple noise-cancellation in vivo Nef staining in B cells analyzed by ImageStreamX MKII. Pre-experimental damaged cells are first excluded with Live/Dead from focused/centroid/singlet image-acquired CD19+ B cells (first lane right/R4 gated on “Focused”). Following Nef+ gating (second lane left/R5 gated on R4), a second step of Nef noise cancellation (second lane middle/R7 on R5) comprises double-negative removal of post-experimental stochastic irregular staining building a disparate intracellular staining gradient (X axis, Nef signal pixel Haralick variance mean) and post-experimental batch overt cellular staining (Y axis, Nef signal pixel intensity threshold). This outputs a B-cell population with a fine-textured pericellular Nefint-lo staining, biologically concordant with Nef membrane-anchoring. Probing of anti-Env BCR (αEnv) by recombinant SIV Env (second lane, right) is combined, resulting in a 2-D panel of intracellular Nef versus αEnv for noise-cancelled Nef+ B cells (R7) plus Nef- B cells (R8) (third lane left/“R7+R8”). pAkt Ser473 expression (third lane right) and cellular morphology (B) was further analyzed. DN, double-negative.

(B) Typical images of Nef-transferred Env-specific B cells defined αEnv+-intracellular Nef+-Nef Haralick variance meanlo-Nef Intensity thresholdlo-Live/Dead--CD19+ cells [“Nef+aEnv+” population of lower left panel in (A), gated on “R7, R8”]. Note the pericellular pAkt Ser473 upregulation in these cells (Ch 03/yellow). Data on inguinal lymph node lymphocytes of macaque R10-007 at week 62 post-SIVmac239 infection are shown in (A) and (B).

(C) Comparison of pAkt Ser473 median fluorescence (medFI) intensity levels in Nef- B cells (R8) versus noise-cancelled Nef+ B cells (R7). Analyzed by paired t test.

(D) Comparison of Nef-positive cell frequencies in non-Env-specific (left) versus Env-specific (right) B cells in lymph nodes of persistently SIV-infected macaques (n = 6). Analyzed by paired t test.

Added in the revised version: Cell contact-dependent B-cell invasion from infected cells in vitro

Added in the revised version: To understand cell-intrinsic properties enhancing B-cell Nef acquisition in vivo, we addressed in an in vitro coculture reconstitution how infected cell-to-B-cell Nef invasion may take place or be modulated. We performed imaging cytometry on a twelve-hour coculture of SIV-disseminating HSC-F cells (reaching around 30% Nef-positivity at moment of coculture) with Ramos B cells (Figure 6A, top left), independent of MHC-related interactions known to enhance T-cell/B-cell contact (Wülfing et al., 1998). Machine learning-verified noise cancellation of fragmented signal acquisition (Figure 6A, middle left) filtered reproducible images of Nef-positive HSC-F cells directly adhering to Ramos B cells within the doublet population (Figure 6A, top to middle right). Importantly, addition of a soluble macaque CD3 stimulation-competent antibody resulted in enhanced Nef+ T-cell adhesion to B cells (Figure 6A, bottom center), accompanying images of polarized Nef protrusion/injection from T cells to B cells (second column, filled wedges) as well as trogocytosed B-cell membranes (third column, open wedge), in keeping with dynamic B-cell Nef acquisition. Centrifugation strengths upon cell collection and staining (190 x g versus 1,200 x g and 600 x g versus 1,200 x g, respectively) shifted the tendency between enriched Nef+ HSC-F-B-cell doublet versus Nef+ B-cell singlet detection via CD3 stimulation, respectively (Figure 6A, lower right).

Added in the revised version: Contact-dependent enhancement of B-cell Nef acquisition in reconstitution.

Added in the revised version: (A) Imaging flow cytometric analyses of stimuli-dependent enhancement of infected HSC-F cell adhesion to non-permissive CD19+ Ramos B cells with the indicated gating (left). Cocultured doublet cells subjected to machine learning-verified filtration of fragmented signal acquisition (via gating Nef pixel signal contrastlo-Nef mean pixel per objecthi cells) deriving typical images of infected HSC-F cell (Nef stain shown in pale blue pseudocolor) adhesion to Ramos cells (CD19 stain shown in purple pseudocolor) are shown. Analyzed by unpaired t test for doublet formation in the unstimulated versus anti-CD3 monoclonal antibody (clone FN-18)-stimulated group.

Added in the revised version: (B) Flow cytometric analyses of Nef-acquiring live CD19+ Ramos B cells upon coculture with SIV-infected HSC-F cells with or without the indicated stimuli. Analyzed by one-way ANOVA with Tukey’s post-hoc multiple comparison tests.

Added in the revised version: (C) Representative flow cytometric plot (left), histogram (middle) and pAkt Ser473 signal deviation in wild type versus G63E Nef-acquiring CD20+ primary B cells (six wells/control). Analyzed by unpaired t test. Data represent pooled data of two (A) or three [(B) and (C)] independent experiments with indicated number of replicates showing similar results.

Added in the revised version: Substantiated by the acquired images, an enhancement of live singlet Ramos B-cell Nef acquisition (Nef+ frequencies/total) was also observed by CD3 stimulation in conventional flow cytometry (Figure 6B). We observed a similar enhancement by phytohemagglutinin (PHA) stimulation whereas phorbol-12-myristate-13-acetate (PMA)/ionomycin did not enhance B-cell Nef acquisition. In a similar 72-hour static infected cell coculture with primary CD20+ B cells, we detected recapitulation of the decreased B-cell pAkt Ser473 phosphorylation upon G63E Nef transfer compared to that of WT Nef (Figure 6C), implicating that the Nef-G63E mutation can directly alleviate maturation inhibition in preferentially Nef-targeted Env-specific B cells.

Removed in the revised version: With this approach we acquired images of a Live/Dead--Nef signal Haralick variance meanlo-Nef signal intensity thresholdlo-Nef+CD19+ B-cells, with fine-textured gradation of Nefint-lo staining with continuity from membrane-proximal regions without sporadic staining speckles (Figure 5A and Figure 5-figure supplement 1A). Images resembled Nef-transferred B cells generated in vitro by a SIV-infected CD4+ T-cell coculture (Figure 5-figure supplement 1B). A linear discriminant analysis-based machine learning module, scoring segregation of typical void/valid images, showed this gating to provide the highest two-dimensional separation (Figure 5-figure supplement 1A). Combined with biotinylated Env/streptavidin staining, this reproducibly obtained images of Env-specific (Ch 12) CD19+ (Ch 11) B cells without membrane ruptures (Ch 08) and showing fine-textured transferred Nef (Ch 02) (Figure 5B). Nef invasion upregulated pAkt Ser473 (Ch 03) to a range resembling in vitro analysis (Figure 5C), stating that Nef-driven aberrant PI3K/mTORC2 signaling does occur in Nef-invaded B cells in vivo and the current phenotypic change likely affects B-cell dynamics. Strikingly, lymph node Env-specific B cells showed significantly higher Nef-positive frequencies as compared with batch non-Env-specific B cells (Figure 5D). This indicated that infected cell-generated Nef preferentially targets adjacent Env-specific B cells, putatively through infected CD4+ cell-to-B-cell contact/transfer (Xu et al., 2009; Hashimoto et al., 2016) and that a Nef phenotypic change likely dictates Env-specific B-cell maturation.

Enhanced Env-specific B-cell responses after PI3K-diminuting mutant selection

Earlier version: Finally, peripheral SIV Env-specific IgG+ B-cell responses comprising plasmablasts (PBs) and memory B cells (Bmem) were analyzed to assess in vivo B-cell quality in NAb inducers with Nef-G63E mutant selection. After excluding lineage-specific cells [T/NK/pro-B cells/monocytes/myeloid dendritic cells (DCs)/plasmacytoid DCs], we defined IgG+ PBs as showing: high replication (Ki-67+), post-activation (HLA-DR+), transcriptional switching for terminal differentiation (IRF4hi) and downregulated antigen surface binding (surface [s]Envlo-cytoplasmic [Cy]Env+) (Nutt et al., 2015; Silveira et al, 2015). Env-specific PBs were defined as CD3-CD8-CD10-CD14-CD16-CD11c-CD123-CD19-CyIgG+Ki-67+IRF4hiHLA-DR+ (total PB) sEnvloCyEnv+ (Env-specific) (Figure 6-figure supplement 1). NAb inducers showed significantly higher Env-specific IgG+ PB responses around week 30 p.i., after Nef-G63E selection, compared to those in NAb non-inducers (Figures 6A and 6B). At 1 year, this difference became enhanced; Env-specific PBs were reduced in non-inducers, while this reduction was ameliorated in the NAb inducers with Nef-G63E.

Revised version: Finally, peripheral SIV Env-specific IgG+ B-cell responses comprising plasmablasts (PBs) and memory B cells (Bmem) were analyzed to assess in vivo B-cell quality in NAb inducers with Nef-G63E mutant selection. After excluding lineage-specific cells [T/NK/pro-B cells/monocytes/myeloid dendritic cells (DCs)/plasmacytoid DCs], we defined IgG+ PBs as showing: high replication (Ki-67+), post-activation (HLA-DR+), transcriptional switching for terminal differentiation (IRF4hi) and downregulated antigen surface binding (surface [s]Envlo-cytoplasmic [Cy]Env+) (Nutt et al., 2015; Silveira et al, 2015). Env-specific PBs were defined as CD3-CD8-CD10-CD14-CD16-CD11c-CD123-CD19-CyIgG+Ki-67+IRF4hiHLA-DR+ (total PB) sEnvloCyEnv+ (Env-specific) (Figure 7A and Figure 7-figure supplement 1). NAb inducers showed significantly higher Env-specific IgG+ PB responses around week 30 p.i., after Nef-G63E selection, compared to those in NAb non-inducers (Figures 7B and 7C). At 1 year, this difference became enhanced; Env-specific PB frequencies were reduced in non-inducers, while this reduction was ameliorated in the NAb inducers with Nef-G63E.

Enhanced SIV Env-specific B-cell output up to NAb induction following Nef-G63E selection.

Added in the revised version: (A) Representative gating (R02-004, week 32 p.i.) of SIV Env gp140-specific memory B cells (Bmem) and plasmablasts (PBs). Two panels for PB gating are shown in lower resolution for visibility. Bmem staining was performed separately and gating are merged with PB panels (second row, panels 2–4). For lineage-specific exclusion (first row, panel 4), CD10 was stained independently for Bmem (see parenthesis).

Earlier version: (A) Changes in SIV Env gp140-specific plasmablast (PB) frequencies in viremic NAb non-inducers (left, n = 12) and Nef-G63E-selecting NAb inducers (right, n = 6). Available samples of twelve viremic NAb non-inducers (including ten with ≥ 1-year survival) and six NAb inducers were tracked. In the right panel, the frequencies in NAb non-inducers are shown in background (gray) for comparison.

Revised version: (B) Changes in SIV Env gp140-specific plasmablast (PB) frequencies in viremic NAb non-inducers (left, n = 12) and Nef-G63E-selecting NAb inducers (right, n = 6). Available samples of twelve viremic NAb non-inducers (including ten with ≥ 1-year survival) and six NAb inducers were tracked. In the right panel, the frequencies in NAb non-inducers are shown in background (gray) for comparison.

Earlier version: (B) Comparison of Env gp140-specific PB frequencies between NAb non-inducers and inducers by Mann Whitney U tests. In the right, the frequencies at NAb induction were compared with those at year 1 in NAb non-inducers with ≥ 1-year survival (n = 10). Bars: mean ± SD.

Revised version: (C) Comparison of Env gp140-specific PB frequencies between NAb non-inducers and inducers by Mann Whitney U tests. In the right, the frequencies at NAb induction were compared with those at year 1 in NAb non-inducers with ≥ 1-year survival (n = 10). Bars: mean ± SD.

Earlier version: (C) Left: vector chart of Env gp140-specific memory B cell (Bmem) and PB levels. Legends for each animal correspond to the ones in (B). Right: NAb non-inducer vectors empirically define a polygonal GC output attractor area Dn (gray area surrounded with dotted lines) on which they converge by and beyond year 1 p.i. The NAb inducer vectors (shown up to the time of NAb induction) remained outside of Dn. At the moment of NAb induction they converged on a second GC output attractor area Di (red area surrounded with dotted lines), mutually exclusive with Dn (P < 0.0001 by Fisher’s exact test on NAb inducer vector convergence frequency within Dn). Legends for year 1 p.i. in the NAb non-inducers and moment of NAb induction in the NAb inducers are specified.

Revised version: (D) Left: vector chart of Env gp140-specific memory B cell (Bmem) and PB levels. Legends for each animal correspond to the ones in (B). Right: NAb non-inducer vectors empirically define a polygonal GC output attractor area Dn (gray area surrounded with dotted lines) on which they converge by and beyond year 1 p.i. The NAb inducer vectors (shown up to the time of NAb induction) remained outside of Dn. At the moment of NAb induction they converged on a second GC output attractor area Di (red area surrounded with dotted lines), mutually exclusive with Dn (P < 0.0001 by Fisher’s exact test on NAb inducer vector convergence frequency within Dn). Legends for year 1 p.i. in the NAb non-inducers and moment of NAb induction in the NAb inducers are specified.

Earlier version: Concomitant determination of Env-specific IgD-CD27+IgG+ Bmem responses allowed pair-wise analysis of Env-specific IgG+ Bmem/PB as a projection of germinal center (GC) output as an indicator of overall B-cell response quality (Zotos et al., 2010) (Figure 6C, left). In this two-dimensional vector charting of Env-specific Bmem/PB responses, vector protrusion towards upper right reasonably represents higher gross GC output of antibody-forming cells (AFCs) (Bmem/PBs). In NAb non-inducers, all vectors converged on an empirically defined polygonal attractor area Dn (gray area surrounded with dotted lines in Figure 6C, right) at year 1 p.i. and beyond, describing that these NAb non-inducers failed to sustain Env-specific GC output. In contrast, the vectors were consistently tracked outside Dn in the NAb inducers with Nef G63E (Figure 6C). At the moment of NAb induction, they converged on another upper-right GC output attractor area Di(red area surrounded with dotted lines in Figure 6C, right), mutually exclusive with Dn (P < 0.0001 by Fisher’s exact test on NAb non-inducer/inducer vector distribution frequency within Dn). These suggest more robust virus-specific IgG+ B-cell responses following Nef-G63E CD8+ T-cell escape mutant selection, predisposing to NAb induction. This enhanced cognate B-cell signature was also an inverted pattern of impaired terminally-differentiated B-cell responses in APDS (Al Qureshah et al., 2021).

Revised version: Concomitant determination of Env-specific IgD-CD27+IgG+ Bmem responses allowed pair-wise analysis of Env-specific IgG+ Bmem/PB as a projection of germinal center (GC) output as an indicator of overall B-cell response quality (Zotos et al., 2010) (Figure 7D, left). In this two-dimensional vector charting of Env-specific Bmem/PB responses, vector protrusion towards upper right reasonably represents higher gross GC output of antibody-forming cells (AFCs) (Bmem/PBs). In NAb non-inducers, all vectors converged on an empirically defined polygonal attractor area Dn (gray area surrounded with dotted lines in Figure 7D, right) at year 1 p.i. and beyond, describing that these NAb non-inducers failed to sustain Env-specific GC output. In contrast, the vectors were consistently tracked outside Dn in the NAb inducers with Nef G63E (Figure 7D). At the moment of NAb induction, they converged on another upper-right GC output attractor area Di (red area surrounded with dotted lines in Figure 7D, right), mutually exclusive with Dn (P < 0.0001 by Fisher’s exact test on NAb non-inducer/inducer vector distribution frequency within Dn). These suggest more robust virus-specific IgG+ B-cell responses following Nef-G63E CD8+ T-cell escape mutant selection, predisposing to NAb induction. This enhanced cognate B-cell signature was also an inverted pattern of impaired terminally-differentiated B-cell responses in APDS (Al Qureshah et al., 2021).

Taken together, in the current model, Nef62-70-specific CD8+ T-cell responses in persistently SIVmac239-infected macaques selected for an escape mutant, Nef-G63E SIV, which decreased aberrant Nef binding and drive of B-cell-inhibitory PI3K/mTORC2. B-cell Nef invasion in vivo occurred more preferentially in Env-specific B cells, suggesting diminution in tonic Nef-mediated B-cell dysregulation after mutant selection. These predisposed to enhanced Env-specific B-cell responses and subsequent SIVmac239-specific NAb induction, altogether, in a manner reciprocal to human APDS-mediated immune dysregulation.

Discussion

Earlier version: In the present study, we found NAb induction against an NAb-resistant SIV after selection of a CD8+ T-cell escape nef mutant virus. The Nef-G63E mutation, positioned in cognate B cells, results in reduction of binding/drive of PI3K/mTORC2, being concordant with in vivo enhancement of Env-specific IgG+ B cells up to NAb induction following mutant selection. Importantly, this manifested through a pAkt deviation level more vivid than wild-type versus germline PIK3CD gain-of-function mutation heterozygote mice showing full APDS phenotype, stating that the current Nef-G63E-associated B-cell/NAb phenotype likely occurs in a “reciprocal APDS-like” manner. This proposition was enhanced based on extension of immune cell-intrinsic Nef influence on cognate B cells (Fig. 5), in addition to infected T cells. This work is, to our knowledge, the first to interlink a PI3K/mTORC2-modulating viral signature and enhanced B-cell/NAb responses in a viral infection model.

Revised version: In the present study, we found NAb induction against an NAb-resistant SIV after selection of a CD8+ T-cell escape nef mutant virus. The Nef-G63E mutation, positioned in cognate B cells, results in reduction of binding/drive of PI3K/mTORC2, being concordant with in vivo enhancement of Env-specific IgG+ B cells up to NAb induction following mutant selection. Importantly, this manifested through a pAkt deviation level more vivid than wild-type versus germline PIK3CD gain-of-function mutation heterozygote mice recapitulating full human APDS phenotype (Avery et al., 2018). Thus, the current Nef-G63E-associated B-cell/NAb phenotype likely occurs in a “reciprocal APDS-like” manner. This proposition was enhanced based on extension of immune cell-intrinsic Nef influence on cognate B cells (Fig. 5), in addition to infected T cells. This work is, to our knowledge, the first to interlink a PI3K/mTORC2-modulating viral signature and enhanced B-cell/NAb responses in a viral infection model.

Earlier version: A link between viral T-cell escape and consequent immune modulation has been previously partially explored, including enhanced mutant HIV-1 epitope peptide binding to inhibitory MHC-I receptor on DCs impairing T cells (Lichterfeld et al., 2007) and decreased CTL-mediated lymph node immunopathology rather driving LCMV-specific antibody production (Battegay et al., 1993). Our results now evidence a new pattern of NAb responses bivalently shaped through viral interactions with both humoral and cellular immunity in AIDS virus infection. SIVmac239 infection of macaques possessing MHC-I alleles associated with Nef-G63E mutation can be one unique model to analyze virus-host interaction for B cell maturation leading to NAb induction. A nef polymorphism was extracted, partially appending to a virological impact of Nef (Kirchhoff et al., 1995; Gauduin et al., 2006; Stolp et al., 2012) in perturbing immunity whereas the current phenotype differed from such replication-related properties. The nef region encoding Nef-G63E is not conserved in HIV-1 (Schindler et al., 2004) and beyond this model study it remains to be addressed whether mutant HIV-1 with a similar phenotype can be obtained.

Revised version: A link between viral T-cell escape and consequent immune modulation has been previously partially explored, including enhanced mutant HIV-1 epitope peptide binding to inhibitory MHC-I receptor on DCs impairing T cells (Lichterfeld et al., 2007) and decreased CTL-mediated lymph node immunopathology rather driving LCMV-specific antibody production (Battegay et al., 1993). Our results now evidence a new pattern of NAb responses bivalently shaped through viral interactions with both humoral and cellular immunity in AIDS virus infection. SIVmac239 infection of macaques possessing MHC-I alleles associated with Nef-G63E mutation can be one unique model to analyze virus-host interaction for B cell maturation leading to NAb induction. The time frame of Nef-G63E mutant selection versus moment of NAb induction was variable in the inducers, for which we did not obtain a clear explanation. This may be influenced by certain basal competitive balance between humoral versus cellular adaptive immunity, as observed in certain vaccination settings (Querec et al., 2009).

Added in the revised version: Identification of the current Nef-G63E candidate phenotype adds one aspect to the wealth of evidence documenting the multifaceted impact of Nef in HIV/SIV infection (Kirchhoff et al., 1995; Gauduin et al., 2006; Stolp et al., 2012). It appeared unique that the current phenotype differed from decreased replication-related properties or generalized amelioration in immune impairment such as those documented in delta-Nef HIV/SIV infection (Kirchhoff et al., 1995; Johnson et al., 1997; Gauduin et al., 2006; Fukazawa et al., 2012). The SIV Nef N-terminal unstructured region comprising Nef-G63E is not conserved in HIV-1 (Schindler et al., 2004) and beyond this model study it remains to be addressed whether a mutant HIV-1 with a similar immunosignaling-affecting phenotype can be obtained and how much lentiviral conservation exists for such interactions. Cognate B-cell maturation analysis on cohort basis (Hau et al., 2022) may potentially assist this approach.

Added in the revised version: Obtained results on preferential Nef transfer to Env-specific B cells (Figure 5) implicate involvement of polarized cognate immune cell perturbation in vivo, best depicted for preferential infectivity of HIV-specific CD4+ T cells (Douek et al., 2002) and HIV condensation to the dendritic cell-T-cell interface (McDonald et al., 2003). As cognate MHC class II interactions (Wülfing et al., 1998; Giltin et al., 2014) regularly occur between HIV/SIV-specific CD4+ T cells and Env-specific B cells in the germinal center, results may show that Env-specific B-cell dysregulation stems from a negative consequence of exploited cognate T-cell/B-cell orchestration. Enhanced T-cell attachment and transfer of Nef to B cells in vitro by CD3 activation (Figure 6), mimicking peptide/MHC (pMHC) signaling, may facilitate recovery from viral T-cell cytoskeletal impairment (Campbell et al., 2004; Jolly et al., 2004; Stolp et al., 2009) antagonizing reduced motility (Stolp et al., 2012), resulting in in vivo “filtering” of enhanced cognate CD4+ T-cell interaction with cognate B cells. Combined with recent in vivo detection of trogocytosed CD20-bearing follicular CD4+ T cells positive for viral DNA and showing increase during SIV infection (Samer et al., 2023), results suggest occurrence of contact-dependent, polarized interactions facilitating Nef transfer to cognate B cells.

Earlier version: The major in vivo readout of this study is autologous neutralization of a highly resistant SIV, differing from HIV-specific broadly neutralizing antibodies (bNAbs). Importantly, however, obtained signatures of enhanced IgG+ Env-specific AFC B cells (class switch/terminal differentiation) and virus neutralization (hypermutation) directly evidence enhanced activity of activation-induced cytidine deaminase (AID), the canonical positive driver of B-cell fate, showing a strong conceptual continuity with bNAb regulation by AID.

Revised version: The major in vivo readout of this study is autologous neutralization of a highly difficult-to-neutralize SIV, differing from HIV-specific broadly neutralizing antibodies (bNAbs). Importantly, however, obtained signatures of enhanced IgG+ Env-specific AFC B cells (class switch/terminal differentiation) and virus neutralization (hypermutation) directly evidence enhanced activity of activation-induced cytidine deaminase (AID), the canonical positive driver of B-cell fate, showing a strong conceptual continuity with bNAb regulation by AID.

Earlier version: Our results extracted an in vivo link between a decreased aberrant Nef-PI3K/mTORC2 interaction and major enhancement in B-cell responses. The relationship is indeed reciprocal to immunogenetic mechanisms of human PI3K gain-of-function mutations resulting in APDS with multiply impaired anti-viral B-cell responses. While the exact hierarchy of binding between Nef and PI3K/mTORC2 components (including PI3K isoforms) and its perturbation by Nef-G63E mutation remains much to be investigated, PI3Kp85 did reasonably show the strongest Nef-binding index in PLA assay (Figure 4D) as Nef and PI3Kp85 possess SH3-binding PxxP and SH3 domains, respectively (Rickles et al., 1994). Binding decrease in canonical partner PI3Kp85 may domino-trigger decrease with others, resulting in quadruple PI3K/mTORC2 binding decrease and fragile downstream Akt signaling.

Revised version: Our results extracted an in vivo link between a decreased aberrant Nef-PI3K/mTORC2 interaction and major enhancement in B-cell responses. The relationship is indeed reciprocal to immunogenetic mechanisms of human PI3K gain-of-function mutations resulting in APDS with multiply impaired anti-viral B-cell responses. APDS occurs from germline mutations in leukocyte-intrinsic p110δ catalytic subunit-coding PIK3CD or the ubiquitous p85α regulatory subunit-coding PIK3R1, both deriving class IA PI3K gain-of-function residue substitutions (Dornan et al., 2017). Such mutations lead to immune cell-intrinsic perturbations with cognate B-cell blockade, manifesting as impaired class-switched B-cell responses and susceptibility against various infections (Angulo et al., 2013; Lucas et al., 2014; Coulter et al., 2017). In the current work, we conversely document a decrease in virally induced tonic PI3K drive, enhanced IgG class-switched cognate B-cell responses and NAb induction against a notably difficult-to-neutralize SIV. While the exact hierarchy of binding between Nef and PI3K/mTORC2 components (including PI3K isoforms) and its perturbation by Nef-G63E mutation remains to be investigated, PI3Kp85 did reasonably show the strongest Nef-binding index in PLA assay (Figure 4D) as Nef and PI3Kp85 possess SH3-binding PxxP and SH3 domains, respectively (Rickles et al., 1994). Binding decrease in canonical partner PI3Kp85 may domino-trigger decrease with others, resulting in quintuple PI3K/mTORC2 binding decrease and fragile downstream Akt signaling. Such domino-triggering is also observed from gain-of-function-mutated p85 to p110 evoking its activation (Dornan et al., 2017) similarly highlighting the fine-tuned nature of the pathway.

Added in the revised version: A limitation of this study is the use of retrospective samples, posing constraints for detecting exact temporal changes in tonic viral B-cell perturbation. Related with this, sampling of lymph node cells in animals belonging to the subgroup of interest at the optimal moment was unpredictable. Certain plasma and cellular samples were also unavailable in differing experiments. While partially addressed in in vitro reconstitution, much remains to be investigated for the observed polarity of Env-specific B-cell perturbation in vivo. Similarly, influence of the Nef-G63E phenotype on CD4+ T cells in vivo and its consequence on B-cell modulation remains largely elusive, and in the current work we focused mainly on molecular properties of the NAb-correlating Nef-G63E mutant strain. As HIV/SIV infection uniquely shows a biphasic Tfh dysregulation of acute-phase destruction (Mattapallil et al., 2005; Moukambi et al., 2015; Moukambi et al., 2019) followed by dysregulated chronic-phase hyper-expansion (Lindqvist et al., 2012; Petrovas et al., 2012; Cubas et al., 2013; Xu et al., 2015), it is of interest how the PI3K drive-ameliorated and proapoptotic mutant influenced this axis post-in vivo selection, warranting further studies ahead. While influence on B cells appears to be stronger by unexplained reasons (Asano et al., 2018), it is noteworthy that PI3K gain-of-function also evokes basal dysregulation in peripheral CD4+ T cells, locally recapitulating the consequence of immunodeficiency virus replication and implicating importance of this signaling axis in HIV/SIV infection. In any event, usage of specific animal models does help to depict a temporal cascade of in vivo events leading to NAb induction, assisting human immunology. As future extension of this study extracting the NAb-involved molecular axis, manipulation/reconstitution experiments shall be designed.

Earlier version: In conclusion, we exemplify in a non-human primate AIDS model that NAb induction against a highly resistant SIV occurs after selection of a CD8+ T-cell escape variant reduced in excess PI3K/mTORC2 drive. These results provide an understanding of how immunodeficiency virus-specific NAb responses can be shaped by both wings of adaptive immune pressure and depicts that immune cell-intrinsic PI3K/mTORC2 manipulation may be pivotal for harnessing antiviral NAb responses. Furthermore, search of other analogies between viral immune dysregulations and human IEIs may lead to identification of novel targets to modulate and harness immune responses in translational settings.

Revised version: In conclusion, we exemplify in a non-human primate AIDS model that NAb induction against a highly resistant SIV occurs after selection of a CD8+ T-cell escape variant reduced in excess PI3K/mTORC2 drive. These results provide an understanding of how immunodeficiency virus-specific NAb responses can be shaped by both wings of adaptive immune pressure and depicts that immune cell-intrinsic PI3K/mTORC2 manipulation may be pivotal for harnessing antiviral NAb responses. As human IEIs are increasingly becoming discovered with various immune-perturbing phenotypes and inheritance patterns (Jauch et al., 2023), search of other types of analogies between viral immune dysregulations and human IEIs may lead to identification of novel targets to modulate and harness immune responses in translational settings.

Removed in the revised version: Limitations of study

Removed in the revised version: A limitation of this study is the use of retrospective samples, posing constraints for detecting exact temporal changes in tonic viral B-cell perturbation. Related with this, sampling of lymph node cells in animals belonging to the subgroup of interest at the optimal moment was unpredictable. Nevertheless, usage of such specific animal models does help to depict a temporal cascade of in vivo events leading to NAb induction, assisting human immunology. As future extension of this study extracting the NAb-involved molecular axis, manipulation/reconstitution experiments shall be designed.

Removed in the revised version: Supporting information